# Numpy Exercises

For this assigment you will need to have numpy installed. 

You can do this by typing 

``pip install -r requirements.txt``

with the `requirements.txt` file in the lecture material, or simply by

``pip install numpy``

or with your favorite package manager.


# Assignment 1: Creating 1D Arrays

Write a function ``assignment_01`` that creates a one dimensional ``numpy`` array with the numbers [ 0,  3,  6,  9, 12, 15, 18, 21, 24, 27, 30].


<!-- grade_id: cell-b2e324c04d6aa864, points: 0, grade: false, solution: true -->
```python
import numpy as np

def assignment_01():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-71d37a088e52639e, points: 1, grade: true, solution: false -->
```python
correct = np.array([0, 3, 6, 9, 12, 15, 18, 21, 24, 27, 30])
your_solution = assignment_01()
assert np.array_equal(your_solution, correct)
```

# Assignment 2: Creating 2D Arrays

Write a function ``assignment_02`` that creates a ``numpy`` array with 4 rows and 5 columns. The content of the array should be 1 for all cells initially. Then, the function should multiply each column by its column index plus 1 and then subtract its row index.  


<!-- grade_id: cell-8553cadd0dee3024, points: 0, grade: false, solution: true -->
```python
import numpy as np

def assignment_02():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5967b51fd21aec5c, points: 4, grade: true, solution: false -->
```python
correct = np.array(
        [
            [1.0, 2.0, 3.0, 4.0, 5.0],
            [0.0, 1.0, 2.0, 3.0, 4.0],
            [-1.0, 0.0, 1.0, 2.0, 3.0],
            [-2.0, -1.0, 0.0, 1.0, 2.0],
        ]
    )
your_solution = assignment_02()
assert np.array_equal(your_solution, correct)
```

# Assignment 3: Indexing 2D Arrays

Write a function ``assignment_03`` that expects a ``numpy`` array with 4 rows and 5 columns. Then, extract the bold-faced subarray and return it

|    |    |    |    |    |
|:-: |:-: |:-: |:-: |:-: |
| .  | .  | .  | .  | .  |
| .  | o  | o  | o  | .  |
| .  | o  | o  | o  | .  |
| .  | .  | .  | .  | .  |


<!-- grade_id: cell-ad862bd904f83cfa, points: 0, grade: false, solution: true -->
```python
import numpy as np

def assignment_03(arr):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f08cd33a2ae12869, points: 2, grade: true, solution: false -->
```python
arr = np.array(
        [
            [1.0, 2.0, 3.0, 4.0, 5.0],
            [0.0, 1.0, 2.0, 3.0, 4.0],
            [-1.0, 0.0, 1.0, 2.0, 3.0],
            [-2.0, -1.0, 0.0, 1.0, 2.0],
        ]
    )
correct = np.array([[1.0, 2.0, 3.0], [0.0, 1.0, 2.0]])
your_solution = assignment_03(arr)
assert np.array_equal(your_solution, correct)
```

# Assignment 4: Manipulating 2D Arrays

Write a function ``assignment_04`` that creates a ``numpy`` array with 100000 rows and 2 columns. The first column should contain Gaussian distributed variables with mean 1 and standard deviation 2 and the second column should contain Gaussian distributed variables with mean -2 and standard deviation 0.5. 



<!-- grade_id: cell-deea6d219694fc91, points: 0, grade: false, solution: true -->
```python
import numpy as np

def assignment_04():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1379a2d9846f9060, points: 2, grade: true, solution: false -->
```python
result = assignment_04()
true_mean = np.array([1, -2])
true_std = np.array([2.0, 0.5])
mean_diff = result.mean(axis=0) - true_mean
std_diff = result.std(axis=0) - true_std
assert all(np.hstack([mean_diff, std_diff]) < 0.1)
```