# ```pandas``` Exercises

For this assigment you will need to have ```Pandas``` installed.

For example, you can install the package by typing

``pip install -r requirements.txt``

with the `requirements.txt` file in the lecture material, or simply by

``pip install numpy``.


## Assignment 01

**Extract some summary statistics of the money spent by the Senat of Berlin**

Write a function ``assignment_01`` that takes the data frame of spendings and returns a list with

- the count
- the mean
- the standard deviation
- the minimum
- the 25% percentile
- the 50% percentile (median)
- the 75% percentile
- the maximum

of all spendings in a list. The data is in the subdirectory ``data`` and can be loaded by ``df = pd.read_csv("data/zuwendungen-berlin.csv.gz")``. For convenient compuation of the summary statistics check the pandas Series API for ``describe()``

```python

def assignment_01(_df):
    spending_statistics = _df. ...
    ...
    return spending_statistics

```


<!-- grade_id: cell-d119af6ffb6a3fca, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

# load the data into a data frame
df = pd.read_csv("data/zuwendungen-berlin.csv.gz")

#print(df.head())
def assignment_01(_df):
    # YOUR CODE HERE
    # stats = _df['Euro'].describe()
    # return stats.tolist()
    
    spendings = _df["Betrag"]
    # get statistics using describe()
    stats = spendings.describe()
    # return as numpy array in correct order
    spending_statistics = np.array([
        stats["count"],
        stats["mean"],
        stats["std"],
        stats["min"],
        stats["25%"],
        stats["50%"],
        stats["75%"],
        stats["max"]
    ])
    return spending_statistics
    #raise NotImplementedError()
```

<!-- grade_id: cell-933dad57e5934437, points: 2, grade: true, solution: false -->
```python
spending_statistics = np.array(
        [
            4.08200000e04,
            2.29215965e05,
            3.93196343e06,
            1.00000000e02,
            4.67300000e03,
            1.64770000e04,
            6.11755000e04,
            4.87261162e08,
        ]
    )
print(assignment_01(df) - spending_statistics)
assert np.allclose(assignment_01(df), spending_statistics)
```

Output:
```
[ 0.00000000e+00 -3.18471284e-05 -1.19002303e-03  0.00000000e+00
  0.00000000e+00  0.00000000e+00  0.00000000e+00  0.00000000e+00]
```

## Assignment 02

**How much is each recipient of a spending receiving in total?**

Write a function ``assignment_02`` that takes the data frame of spendings and groups by recipient (column ``'Name'``) and then sums all money received for each recipient. Return the names of the recipients that received in total 143 Euros. 

```python

def assignment_02(df):
    money_received = df.groupby(['Name']). ...
    ...
    return names_of_recipients

```

<!-- grade_id: cell-7082d9f964916a7b, points: 0, grade: false, solution: true -->
```python
def assignment_02(df):
    # YOUR CODE HERE
    # group by recipient name and sum all spending amounts
    money_received = df.groupby("Name")["Betrag"].sum()
    # find all recipients whose total is exactly 143
    names_of_recipients = money_received[money_received == 143].index.tolist()
    return names_of_recipients
    #raise NotImplementedError()
```

<!-- grade_id: cell-643077937c9581fa, points: 3, grade: true, solution: false -->
```python
result = sorted(assignment_02(df))
assert (result[0] == "Rock 'n' Roll Club Pinguin Berlin e. V.") & (
    result[1] == "Triathlongemeinschaft Sisu Berlin e. V."
)
```

## Assignment 03

**How much is Berlin spending on each political ressort?**

Write a function ``assignment_03`` that takes the data frame of spendings (spending is the column 'Betrag'), groups by political ressort (in german 'Politikbereich') and computes the 

 - minimum
 - median
 - maximum

of the spendings on each political ressort. Return the aggregates in the political ressort ('Politikbereich') 'sciences' ('Wissenschaft')

```python

def assignment_03(df):
    spending_per_ressort = df.groupby(['Politikbereich']). ...
    ...
    return 

```


<!-- grade_id: cell-3d2ea05dc8e37bd4, points: 0, grade: false, solution: true -->
```python
def assignment_03(df):
    # YOUR CODE HERE
    
    # group by political ressort and calculate min, median and max
    spending_per_ressort = df.groupby("Politikbereich")["Betrag"].agg(
        ["min", "median", "max"]
    )
    result = spending_per_ressort.loc["Wissenschaft"]
    # return as numpy array in correct order
    return np.array([result["min"], result["median"], result["max"]])
    
    #raise NotImplementedError()
```

<!-- grade_id: cell-25519449d6309967, points: 3, grade: true, solution: false -->
```python
correct = np.array([500.0, 115557.5, 41852102.0])
assert np.array_equal(assignment_03(df), correct)
```

## Assignment 04

**How much is Berlin spending on each U-Bahn?**

Write a function ``assignment_04`` that takes the data frame of spendings, filters for transportation (german 'Verkehr'), groups by the specific ubahn and sums up the spendings. For the ubahn grouping you can extract the ubahn with the regular expression ``'U[1-9]'``. The function should return the ubahn names ordered from most (first element) to least expensive (last element).


```python

def assignment_04(df):
    df['ubahn'] = df['Zweck'].str.extract('(U[1-9])') ...
    ...
    return 

```


<!-- grade_id: cell-ea979368626cbf9a, points: 0, grade: false, solution: true -->
```python
def assignment_04(df):
    # YOUR CODE HERE
    transport = df[df["Politikbereich"] == "Verkehr"].copy() 
    # extract U-Bahn line from the column 'Zweck'
    transport["ubahn"] = transport["Zweck"].str.extract(r"(U[1-9])")
    # drop rows where no U-Bahn was found
    transport = transport.dropna(subset=["ubahn"])
    # group by U-Bahn line and sum spendings
    spending_per_ubahn = transport.groupby("ubahn")["Betrag"].sum()
    # sort from highest to lowest spending
    sorted_ubahn = spending_per_ubahn.sort_values(ascending=False)
    return sorted_ubahn.index.tolist()
    
    #raise NotImplementedError()
```

<!-- grade_id: cell-3ec58620457dfb69, points: 3, grade: true, solution: false -->
```python
ubahn_cost_ranking = ["U5", "U2", "U1", "U6", "U8", "U7", "U9", "U3", "U4"]
assert all([x == y for x, y in zip(assignment_04(df), ubahn_cost_ranking)])
```