# Assignment 1: Functions (2 Points)

- Write a function ``assignment_01`` that multiplies an arbitrarily long list of number arguments (not passed as a list/collection but as single arguments)
- Add some documentation


<!-- grade_id: cell-a0f77fb0575c9fde, points: 0, grade: false, solution: true -->
```python
def assignment_01(*args):
    product = 1
    # YOUR CODE HERE
    """
    Multiplying all numbers passed as arguments.
    """
    for n in args:
        product = product * n
    #raise NotImplementedError()
    return product
```

<!-- grade_id: cell-5da53998f21b8662, points: 2, grade: true, solution: false -->
```python
assert assignment_01(1,2,3)==6
assert assignment_01(1,2,3,0)==0
```

# Assignment 2a: Loops (2 Points)

- Write a function ``assignment_02_02a`` that expects a list with numbers
- The function should initialize an empty list ``result`` and then iterate though the input list with a for loop
- For each element of the input list, the function should add a string "higher" if the element was larger than 5 and "lower" otherwise
- Then the list ``result`` should be returned by that function


<!-- grade_id: cell-096eddaaab6eb060, points: 0, grade: false, solution: true -->
```python
def assignment_02a(numbers):
    result = []
    for number in numbers:
    # YOUR CODE HERE
        if number > 5:
            result.append('higher')
        else:
            result.append('lower')
    #raise NotImplementedError()
    return result
```

<!-- grade_id: cell-107793b842ff4591, points: 2, grade: true, solution: false -->
```python
numbers = [1,6]
assert assignment_02a(numbers)==['lower','higher']
```

# Assignment 2b: List comprehensions (2 Points)


- Now write a function ``assignment_02b`` that does the same in one line with list-comprehension


<!-- grade_id: cell-77420a1acc200c98, points: 0, grade: false, solution: true -->
```python
def assignment_02b(numbers):
    # YOUR CODE HERE
    return ['higher' if number > 5 else 'lower' for number in numbers]
    #raise NotImplementedError()
```

<!-- grade_id: cell-6469dc1368d2cb68, points: 2, grade: true, solution: false -->
```python
numbers = [1,6]
assert assignment_02b(numbers)==['lower','higher']
```

# Assignment 3: Classes, Functions and Errors (4 Points)
- Write a Class ``Duck``
- A ``Duck`` constructor should accept variable length string arguments denoting the names of the duck
- Internally the names are joined with a " " separator and stored as a single string
- A ``Duck`` has a function ``talks_to`` that 
 - accepts one input argument 
 - returns a string ``Quak`` if the input was another duck and raises a Value Error with the message "Can only talk to other Ducks" otherwise
 - **Optional**: Create a function ``__repr__`` that prints the name of the Duck


<!-- grade_id: cell-2b333fe63d33fed1, points: 0, grade: false, solution: true -->
```python
class Duck:
    # YOUR CODE HERE
    def __init__(self, *names):
        self.name = " ".join(names)
    
    def talks_to(self, other):
        if isinstance(other, Duck):
            return "Quak"
        else:
            raise ValueError("Can only talk to other Ducks")
            
    def __repr__(self):
        return self.name 
    #raise NotImplementedError()
```

<!-- grade_id: cell-41da6ed94694a72a, points: 4, grade: true, solution: false -->
```python
duck = Duck("Alfred", "Judokus", "Quak")
assert duck.name == 'Alfred Judokus Quak'

other_duck = Duck("Hans", "Peter", "Schnabel")
assert other_duck.name == 'Hans Peter Schnabel'
assert duck.talks_to(other_duck) == "Quak"

try:
    duck.talks_to("Crocodile")
except ValueError as ve:
    assert str(ve) == "Can only talk to other Ducks"
```